A unique cell division protein critical for the assembly of the bacterial divisome

Author response:

The following is the authors’ response to the previous reviews.

(1) The reviewers asked to clarify the BTH assay: The fused T25 and T18 domains must be in the cytoplasmic to complement successfully. The authors stated that the N terminus of Aeg1 transverses the membrane once, which means that the T25-Aeg1 will have T25 in the periplasm. However, T18C vector fusion with other division proteins will have T18C of ZipA in the periplasm (ZipA's N terminus is on the periplasmic side of the inner membrane) while that of FtsN in the cytoplasm (FtsN's N terminus is in the cytoplasm). As such, it isn't easy to understand why T25-Aeg1 showed positive results for both ZipA and FtsN. Note that FtsL, FtsB, and FtsI all have the same topology as FtsN but showed negative results. It is possible that these fusion proteins do not fold correctly, and hence, the results cannot be interpreted directly. The authors did not address this concern but only cited that BTH is a commonly used assay for protein-protein interactions.

In response to the editor's comments and the concerns raised by the reviewer, we have performed two sets of Aeg1-T25 fusion experiments to determine whether the Aeg1 topology impacts protein interactions measured by bacterial two-hybrid (BTH) assays. In the first set of experiments, we fused the T25 domain to the N-terminus of Aeg1 and still observed strong binding of Aeg1 to ZipA and FtsN, respectively. Similar results were obtained from the second set of experiments in which the T25 domain was fused to the C-terminus of Aeg1.

These results indicate that the precise topology of Aeg1 does not significantly impact its ability to engage these binding partners. Aeg1 is predicted to harbor a single transmembrane domain, however, the precise location of this transmembrane segment differs in predictions made by different algorithms. The SMART Web site (1) predicted the transmembrane region to be located at the N-terminus of Aeg1 (7-29 aa). In contrast, Phobius, based on HMM （2， 3）suggested the transmembrane segment is situated more centrally within the Aeg1 protein (134-151 aa), and further proposed that the N-terminus may function as a signal peptide. This latter prediction also provides a potential explanation for the larger-than-expected molecular weight of the Aeg1 truncation mutant observed in the Western blot shown in Fig 1C. The removal of the putative signal peptide may have altered the protein structure, affecting its electrophoretic mobility. As a result, we are more inclined to favor the topology model for Aeg1 predicted by Phobius.

(2) It is still difficult to identify the midcell localization patterns of Aeg1 and other division proteins from microscopy images (Fig. 4C and Fig. 5A). In Fig 4C, only ZipA and Aeg1 formed clear, regular band-like colocalization patterns. Others formed irregular co-localized puncta along the cell length, different from the expected midcell localization patterns. Cells also appeared to be much longer than WT cells, suggesting cell division defects. The most likely reason for these aberrant localization patterns and filamentous cells is that GFP/mCherry-fusions of these division proteins are not functional and become dominant negative, interfering with proper cell division. The authors need to test the functionality of these fusion proteins before they can be used for imaging. (The authors also mislabeled Hoechst and the division protein GFP panels labels in this figure.)

Thank you for raising this important point. To examine the functionality of the fluorescence protein fusion constructs, we have painstakingly performed conditional knockout of the genes of interest (zipA, ftsB, ftsL, and ftsN) in A. baumannii strains inducibly expressing the corresponding fusion protein. We found that these fluorescence protein fusions were able to fully rescue the growth of the mutant lacking the corresponding fts gene (Figure 4-figure supplement 1). Concurrently, we have also successfully knocked out the aeg1 gene under conditions in trans expression of an mCherry-Aeg1 fusion protein, which was able to effectively rescue the growth defects of the Δa_eg1_ mutant (Figure 4-figure supplement 1). We then introduced the functional fluorescence protein fusions into wild-type cells and observed the co-localization of Aeg1 with the relevant Fts proteins. The results showed that Aeg1 indeed co-localized with ZipA, FtsB, FtsL, and FtsN (Fig.4E, red arrows), but occasional non-co-localization was also observed (Fig.4E, white arrows).

We have utilized the functional fluorescence protein fusion constructs to analyze the localization of relevant Aeg1-interacting proteins in the Δaeg1 strain upon Aeg1 depletion. Our results showed that the depletion of Aeg1 indeed impacted the midcell localization of the several Aeg1-interacting Fts proteins.

References

(1) Letunic I, Khedkar S, Bork P. SMART: recent updates, new developments and status in 2020. Nucleic acids research. 2021;49:D458-d60.doi: 10.1093/nar/gkaa937.

(2) Käll L, Krogh A, Sonnhammer EL. A combined transmembrane topology and signal peptide prediction method. Journal of molecular biology. 2004;338:1027-36.doi: 10.1016/j.jmb.2004.03.016

(3) Käll L, Krogh A, Sonnhammer EL. Advantages of combined transmembrane topology and signal peptide prediction--the Phobius web server. Nucleic acids research. 2007;35:W429-32.doi: 10.1093/nar/gkm256

Changes from the earlier version to the revised version

Abstract

Abstract

Identification of unique essential bacterial genes is important for not only the understanding of their cell biology but also the development of new antimicrobials. Here we report a previously unrecognized core component of the Acinetobacter baumannii divisome. Our results reveal that the protein, termed Aeg1 interacts with multiple cell division proteins, including FtsN, which is required for components of the divisome to localize to the midcell. We demonstrate that the FtsAE202K and FtsBE65A mutants effectively bypassed the need of Aeg1 by A. baumannii, as did the activation variants FtsWM254I and FtsWS274G. Our results suggest that Aeg1 is a cell division protein that arrives at the division site to initiate cell division by recruiting FtsN, which activates FtsQLB and FtsA to induce the septal peptidoglycan synthase FtsWI. The discovery of the new essential cell division protein has provided a new target for the development of antibacterial agents.

Introduction

Acinetobacter baumannii is a Gram-negative bacterial pathogen that causes a wide variety of nosocomial infections (1, 2), including urinary tract infections, gastrointestinal and skin/wound infections and secondary meningitis (3–6). The challenge posed by this pathogen is compounded by the emergence of isolates that are resistant to almost all of the available antibiotics, including some of the last resort for treatment of infections caused by bacterial pathogens, including colistin, tigecycline, and carbapenems (2, 7). As a result, the World Health Organization (WHO) has classified A. baumannii as a pathogen for which new antibiotics are urgently needed (8). Meeting this challenge requires the discovery of antimicrobials of new mechanisms with action or other mitigation strategies that enable reduction of disease burden caused by this pathogen.

Historically, the discovery of antibiotics was largely based on the search for secondary metabolites produced by soil-dwelling microbes (9, 10). The majority of clinically used antibiotic classes originate from soil bacteria, including lactams, aminoglycosides, polymyxins, and glycopeptides (11). Unfortunately, natural products discovery is now plagued by problems of so-called ‘dereplication problem’, wherein rediscovery of molecules or their close relatives (12). As a result, there has not been a novel antibiotic discovered in over 30 years using traditional methods of screening environmental samples (13). The overuse of these merely variants of existing compounds has further accelerated the problem associated with antibiotic resistance. Thus, identification of novel therapeutic targets in the pathogen may serve as an effective alternative to the problem.

A better understanding of gene function associated with A. baumannii pathogenesis and biology will provide us strategies to develop antibiotics with potentially novel mechanisms of action. One extensively explored avenue in this regard is the identification of previously unrecognized essential genes by methods enabled by new DNA sequencing technologies, including insertion sequencing (INSeq) (14), transposon sequencing (Tn-seq) (15), transposon-directed insertion site sequencing (TraDIS) (16), and high-throughput insertion tracking sequencing (HITS) (17). Application of these methods has identified genes believed to be required for bacterial viability under various conditions in such important pathogens as Streptococcus pneumoniae (15), Pseudomonas aeruginosa (18), Vibrio cholerae (19) and A. baumannii (20–22).

Similar to candidate genes identified by other methods, putative essential genes obtained by In-seq were deduced from sequencing results of surviving mutants, some of which can be false positives or negatives due to unsaturated mutant libraries or polar effects on downstream genes of the insertion sites. Using In-seq Wang et al identified 453 candidate essential genes for the growth of A. baumannii strain ATCC 17978 on rich medium (22). Among these, genes annotated as hypotheticals with orthologs in other taxonomically distant microorganisms are of particular interest because their potential role in cellular processes that can be explored as targets for the development of novel antimicrobials. We thus set out to test a subgroup of the predicted essential genes by a conditional gene deletion strategy based on the arabinose inducible promoter (23). From ten genes examined, three were found to be essential for A. baumannii viability on rich medium. Further study revealed that one of these genes codes for a novel bacterial cell division protein.

Results

Identification of genes required for A. baumannii viability

Analysis of strains of A. baumannii, including the model strain 17978 by In-seq technologies have led to the identification of at least 453 genes potentially required for bacterial growth on rich medium (21, 22). Whereas many of these genes are involved in known cellular processes fundamental for cell viability such as DNA replication, transcription and translation, some code for proteins of unknown activity, are thus of great interest as potentially new antimicrobial targets (21). To determine the function of the candidate essential genes annotated as hypotheticals, we first confirmed their essentiality by a conditional gene deletion method. To this end, we employed a set of plasmids that allow inducible gene expression in A. baumannii (24). Each gene of interest was expressed from the arabinose-inducible promoter (23), and the corresponding gene was deleted from the chromosome by the standard allele exchange method (25) (Fig. 1A). Chromosomal deletion mutants harboring the arabinose-inducible gene were tested for essentiality by spotting serially diluted cells onto solid medium with or without arabinose (ara) (Fig. 1A). If the gene of interest is essential, the bacterium should fail to grow on plates without ara.

Identification of aeg1 as an essential gene for A. baumannii viability on rich medium

A. A diagram depicting the method for conditional deletion of potential essential genes. The gene to be examined was expressed from the arabinose-inducible promoter on a plasmid and the gene was deleted from the chromosome using the standard allele exchange method (upper panel) and strains lacking the chromosomal gene were tested for essentiality by spotting diluted cells on media with and without arabinose, respectively (lower panel).

B. Identification of A1S_0137, A1S_2249, and A1S_3387 as A. baumannii essential genes. Each of the 10 genes predicted to be essential were deleted by the method described in A and the resulting bacterial strains were tested for growth by spotting serially diluted cells on medium with or without 1% arabinose. Images were acquired after incubation at 37°C for 18 h. Similar results were obtained in at least three independent experiments.

Earlier version: C-D. The predicted transmembrane (TM) domain in the amino terminal end of Aeg1 is required for its activity. A diagram illustrating the location of the TM domain which is predicted to span between the 7th and the 29th residues of Aeg1 (C). Cells harboring differentially regulated HA-Aeg1 and HA-Aeg1ΔTM spotted onto LB agar supplemented with 0.5 mM IPTG or 1% ara, images were acquired after 18 h incubation at 37°C (D upper two panels). The expression of the Aeg1 and Aeg1ΔTM were examined by immunoblotting with the HA-specific antibody. The metabolic enzyme isocitrate dehydrogenase (ICDH) was probed as a loading control (D, lower two panels). Similar results were obtained in three independent experiments.

Revised version: C. Cells harboring differentially regulated HA-Aeg1 and HA-Aeg1ΔN spotted onto LB agar supplemented with 0.5 mM IPTG or 1% ara, images were acquired after 18 h incubation at 37°C (C upper two panels). The expression of the Aeg1 and Aeg1ΔN were examined by immunoblotting with the HA-specific antibody. The metabolic enzyme isocitrate dehydrogenase (ICDH) was probed as a loading control (C, lower two panels). Similar results were obtained in three independent experiments.

Earlier version: Our initial analysis of 10 genes (Table S1) identified A1S_0137, A1S_2249, and A1S_3387 (26) that met the criteria of being required for growth on LB agar (Fig. 1B). Importantly, gene A1S_3387 is also essential for A. baumannii growth in the Vogel-Bonner minimal medium with succinate (VBS) as the sole carbon source (27) (Fig. S1). These genes were thus designated as A. baumannii essential genes (Aeg) and Aeg1 (A1S_3387) was chosen for detailed study.

Revised version: Our initial analysis of 10 genes (Supplementary Files 1) identified A1S_0137, A1S_2249, and A1S_3387 (26) that met the criteria of being required for growth on LB agar (Fig. 1B). Importantly, gene A1S_3387 is also essential for A. baumannii growth in the Vogel-Bonner minimal medium with succinate (VBS) as the sole carbon source (27) (Fig.1-figure supplement 1). These genes were thus designated as A. baumannii essential genes (Aeg) and Aeg1 (A1S_3387) was chosen for detailed study.

Earlier version: Sequence analysis revealed that Aeg1 has orthologs in a number of Gram-negative bacteria, including Klebsiella pneumoniae, Salmonella enterica members of the Moraxellaceae family and some strains of Escherichia coli, with sequence similarity ranging from 80% to identical (Fig. S2). Aeg1 is predicted to harbor a single transmembrane domain composing of 23 residues close to its amino terminal end (Fig. 1C). To determine the importance of the putative membrane spanning region in its function, we employed a dual inducible system with two compatible plasmids that use promoters that respond to ara and isopropyl-β-D-thiogalactopyranoside (IPTG), respectively. Strain Δaeg1(PBAD::Aeg1, PTAC::Aeg1) was able to grow on medium containing 1% ara or 0.5 mM IPTG (Fig. 1D) indicating that the chromosomal deletion can be complemented by either of the constructs when expression of Aeg1 was induced. Aeg1 is a bitopic integral membrane protein similar to ZipA (28). It has a similar arrangement where the N terminus of the protein traverses the membrane once, while the rest of the protein present in the cytoplasm. Importantly, expression of the Aeg1 mutant lacking the predicted membrane anchoring region (residues 7-29) was unable to restore the growth of the strain (Fig. 1D), suggesting that association with membrane is essential for its activity.

Revised version: Sequence analysis revealed that Aeg1 has orthologs in a number of Gram-negative bacteria, including Klebsiella pneumoniae, Salmonella enterica members of the Moraxellaceae family and some strains of Escherichia coli, with sequence similarity ranging from 80% to identical (Fig. 1-figure supplement 2). Aeg1 is predicted to harbor a single transmembrane domain, however, the precise location of this transmembrane segment has been predicted differently across various analyses. SMART Web site (28) predicted the transmembrane region to be located at the N-terminus of Aeg1 (7-29 aa). In contrast, Phobius, based on HMM suggested the transmembrane segment is situated more centrally within the Aeg1 protein (134-151 aa), and further proposed that the N-terminus may correspond to a signal peptide (29, 30). Regardless of the specific transmembrane topology prediction, the N-terminal region of Aeg1 appears to be a critical domain. Therefore, to further elucidate the structural and functional importance of the Aeg1 close to its N-terminus end, we employed a dual inducible system with two compatible plasmids that use promoters that respond to ara and isopropyl-β-D-thiogalactopyranoside (IPTG), respectively. Strain Δaeg1(PBAD::Aeg1, PTAC::Aeg1) was able to grow on medium containing 1% ara or 0.5 mM IPTG (Fig. 1C) indicating that the chromosomal deletion can be complemented by either of the constructs when expression of Aeg1 was induced. Importantly, expression of the Aeg1 mutant lacking N-terminal region was unable to restore the growth of the strain (Fig. 1C), suggesting that N-terminal domain of Aeg1 is essential for its activity.

Depletion of Aeg1 caused cell filamentation in A. baumannii

Earlier version: We next examined the morphology of A. baumannii lacking Aeg1 by optical and scanning electron microscopy. Cultures of strain Δaeg1(PBAD::Aeg1) were induced by ara for 14 h, cells diluted into fresh LB with or without ara were allowed to grow for 18 h. We found that cells of cultures grown in medium containing 1% ara exhibited a short rod cell shape similar to that of the wild-type strain (Fig. 2A). In medium without ara most cells of strain Δaeg1(PBAD::Aeg1) displayed an elongated morphology which is in sharp contrast to those of wild-type cells grew under identical conditions (Fig. 2A). Quantitative analysis indicated that in medium supplemented with ara, more than 95% of the cells of strain Δaeg1(PBAD::Aeg1) were shorter than 5 μm. Incubation in medium without ara for 18 h caused more than 96% of the cells to become elongated with cell lengths ranging from 5 μm to over 10 μm (Fig. 2B). The phenotype can be partially reversed after growing the elongated cells in medium containing 1% ara for 4 h (Fig. 2C). These results indicate that depletion of Aeg1 causes cell elongation, a phenotype that had been observed in recent studies using transposon insertion sequencing (20, 29).

Revised version: We next examined the morphology of A. baumannii lacking Aeg1 by optical and scanning electron microscopy. Cultures of strain Δaeg1(PBAD::Aeg1) were induced by ara for 14 h, cells diluted into fresh LB with or without ara were allowed to grow for 18 h. We found that cells of cultures grown in medium containing 1% ara exhibited a short rod cell shape similar to that of the wild-type strain (Fig. 2A). In medium without ara most cells of strain Δaeg1(PBAD::Aeg1) displayed an elongated morphology which is in sharp contrast to those of wild-type cells grew under identical conditions (Fig. 2A). Quantitative analysis indicated that in medium supplemented with ara, more than 95% of the cells of strain Δaeg1(PBAD::Aeg1) were shorter than 5 μm. Incubation in medium without ara for 18 h caused more than 96% of the cells to become elongated with cell lengths ranging from 5 μm to over 10 μm (Fig. 2B). The phenotype can be partially reversed after growing the elongated cells in medium containing 1% ara for 4 h (Fig. 2C). These results indicate that depletion of Aeg1 causes cell elongation, a phenotype that had been observed in recent studies using transposon insertion sequencing (20, 31).

Aeg1 deficiency causes cell elongation in A. baumannii.

A. Cell morphology of the wild-type and strain Δaeg1(PBAD::Aeg1) grown in medium with or without ara. The Δaeg1(PBAD::Aeg1) and the wild-type strains were grown in medium with or without ara for 18 h, samples were processed for scanning electron microscope imaging. Images shown are representative of three parallel cultures. Bar, 2 µm.

B-C. Expression of Aeg1 reversed the cell elongation phenotype caused by its depletion. Saturated bacterial cultures of Δaeg1(PBAD::Aeg1) diluted in fresh LB broth with (I) or without (II) ara were incubated for 18 h. The culture from the uninduced sample was split into two subcultures, ara was added to one of the subcultures and the morphology of the cells was determined at 4 h (III) post induction. Images are representative of three parallel cultures. Bar, 10 µm (B). Quantitation of the cell length in samples differently expressing Aeg1. Cell length was categorized into groups of 0-5 μm, 5-10 μm, and more than 10 μm. At least 300 cells were counted for each sample (C). Results shown were from the average of three independent experiments. Statistical analysis in each panel was performed by Student’s t test.

Mutations in the cell division gene ftsA bypasses the requirement of Aeg1

Earlier version: To explore the cellular process participated by Aeg1, we sought to isolate suppressor mutants that can survive in standard medium under conditions the gene is not expressed. To this end, cells of strain Δaeg1(PBAD::Aeg1) grown in medium supplemented with ara were exposed to the mutagen ethyl methanesulfonate (EMS), and the treated cells were plated on LB agar without ara. A few colonies were obtained from approximately 6×108 cells and two independent mutants M1 and M2 were chosen for further analysis. In spotting assays using diluted cells, each of these two mutants derived from strain Δaeg1(PBAD::Aeg1) was able to grow on medium without ara at rates comparable to those of wild-type bacteria (Fig. 3A). We then performed whole genome sequencing analysis of these two mutants, which revealed that M1 and M2 had 320 and 328 substitution mutations, respectively. In addition, 659 insertion-deletion (InDel) mutations were found in M1 and 665 such mutations were associated with M2 (Table S2).

Revised version: To explore the cellular process participated by Aeg1, we sought to isolate suppressor mutants that can survive in standard medium under conditions the gene is not expressed. To this end, cells of strain Δaeg1(PBAD::Aeg1) grown in medium supplemented with ara were exposed to the mutagen ethyl methanesulfonate (EMS), and the treated cells were plated on LB agar without ara. A few colonies were obtained from approximately 6×108 cells and two independent mutants M1 and M2 were chosen for further analysis. In spotting assays using diluted cells, each of these two mutants derived from strain Δaeg1(PBAD::Aeg1) was able to grow on medium without ara at rates comparable to those of wild-type bacteria (Fig. 3A). We then performed whole genome sequencing analysis of these two mutants, which revealed that M1 and M2 had 320 and 328 substitution mutations, respectively. In addition, 659 insertion-deletion (InDel) mutations were found in M1 and 665 such mutations were associated with M2 (Supplementary Files 2).

Mutations in proteins involved in cell division bypass the requirement of Aeg1.

A. Isolation of two suppressor mutants that gained viability in the absence of Aeg1. Cells serially diluted with water were spotted onto LB agar with or without ara. Note that the Δaeg1(PBAD::Aeg1) strain (top row) cannot grow on medium without ara and the wild-type strain (2nd row) can grow on both conditions. Similar results were obtained in three independent experiments.

B. The FtsAE202K mutant suppressed the requirement of Aeg1. Cells of strains derived from Δaeg1(PBAD::aeg1) that harbored the vector (I), pHA-Aeg1 (II), pHA-FtsA (III), or pHA-FtsAE202K (IV) were spotted onto LB agar containing ara (left) or IPTG (right), Images were acquired after 18 h incubation at 37°C (b, upper panels). Expression of Aeg1, FtsA or FtsAE202K in these strains were detected by immunoblotting with the HA-specific antibody after lysates of IPTG-induced cells being resolved by SDS-PAGE. ICDH was probed as a loading control (b, lower panels). Similar results were obtained in three independent experiments.

C. The suppression mutants assumed normal cell morphology. Bacterial strains derived from Δaeg1(PBAD::aeg1) grown in LB broth containing ara for 6 h were diluted into fresh medium with the inducer and the cultures were induced with IPTG for 4 h prior to being processed for imaging. Images were representatives of three parallel cultures. Bar, 10 µm.

D. Additional FtsA mutants bypassed the need of Aeg1. Mutant FtsAD124A, FtsAV144A, or FtsAQ285A expressed from the IPTG-inducible PTAC was introduced into strain Δaeg1(PBAD::Aeg1) and serially diluted cells of resulting strains were spotted onto LB agar supplemented with ara or IPTG. Images were acquired after incubation at 37°C for 18 h (upper panels). Expression of Aeg1, FtsAD124A, FtsAV144A, and FtsAQ285A in these strains. Total protein of IPTG-induced cells resolved by SDS-PAGE and proteins transferred onto nitrocellulose membranes were detected by immunoblotting with the HA-specific antibody. ICDH was probed as a loading control (lower panels). Similar results were obtained in three independent experiments.

E. The Aeg1-mCherry fusion localizes to division constrictions. Wild-type A. baumannii harboring PBAD::Aeg1-mCherry grown to the mid-log phase in LB broth containing 1% ara were processed for imaging. Bar, 10 µm. Images are representative of three parallel cultures.

Earlier version: Because depletion of Aeg1 caused cell elongation, we first examined mutated genes whose functions are known to be involved in cell division. An E202K mutation caused by a G to A substitution in the first position of the codon was found in the cell division protein FtsA in both M1 and M2 (Table S2). We thus verified the suppression phenotype by introducing a plasmid expressing FtsA or FtsAE202K into strain Δaeg1(PBAD::Aeg1) and examined the growth of the resulting bacterial strains under different induction conditions. Expression of FtsAE202K allowed strain Δaeg1(PBAD::Aeg1) to gain the ability to grow in the absence of ara. In contrast, although the protein was expressed at a comparable level, FtsA cannot suppress the growth defect caused by Aeg1 depletion (Fig. 3B). Consistent with the growth phenotypes, expression of FtsAE202K in strain Δaeg1(PBAD::Aeg1) under conditions in which the complementation gene was not expressed displayed a morphology similar to that of wild-type bacteria (Fig. 3C).

Revised version: Because depletion of Aeg1 caused cell elongation, we first examined mutated genes whose functions are known to be involved in cell division. An E202K mutation caused by a G to A substitution in the first position of the codon was found in the cell division protein FtsA in both M1 and M2 (Supplementary Files 2). We thus verified the suppression phenotype by introducing a plasmid expressing FtsA or FtsAE202K into strain Δaeg1(PBAD::Aeg1) and examined the growth of the resulting bacterial strains under different induction conditions. Expression of FtsAE202K allowed strain Δaeg1(PBAD::Aeg1) to gain the ability to grow in the absence of ara. In contrast, although the protein was expressed at a comparable level, FtsA cannot suppress the growth defect caused by Aeg1 depletion (Fig. 3B). Consistent with the growth phenotypes, expression of FtsAE202K in strain Δaeg1(PBAD::Aeg1) under conditions in which the complementation gene was not expressed displayed a morphology similar to that of wild-type bacteria (Fig. 3C).

Earlier version: Previous studies have found that several mutations in FtsA bypassed the requirement of a number of essential cell division proteins in E. coli, including ZipA (30), FtsK (31), and FtsN (32). Among these, FtsAR286W (FtsA*) was the first gain-of-function mutation found in E. coli, which allowed the survival of mutants lacking ZipA (ΔzipA) or FtsK (ΔftsK) (30, 31, 33). Additionally, FtsAI143L(34) and FtsAE124A(32) weakly bypassed the requirement of FtsN. These mutations also rescued the FtsLL86F/E87K mutant when overexpressed (35). We thus examined whether similar mutations in FtsA of A. baumannii can bypass the need of Aeg1 by constructing substitution FtsA mutants D124A, V144L and Q285W, which correspond to E124A, I143L, and R286W in its E. coli counterpart, respectively (Fig. S3) and tested them in strain Δaeg1(PBAD::Aeg1). Whereas FtsAQ285W cannot detectably suppress Aeg1 depletion, expression of FtsAD124A or FtsAV144L partially resecured the growth of strain Δaeg1(PBAD::Aeg1) in medium without ara (Fig. 3D). Despite multiple attempts, we were unable to detect the expression of these FtsA mutants by immunoblotting using antibody specific to the Flag epitope fused to these alleles (Fig. 3D, lower panels), suggesting that the low activity exhibited by these mutants may result from low protein levels.

Revised version: Previous studies have found that several mutations in FtsA bypassed the requirement of a number of essential cell division proteins in E. coli, including ZipA (32), FtsK (33), and FtsN (34). Among these, FtsAR286W (FtsA*) was the first gain-of-function mutation found in E. coli, which allowed the survival of mutants lacking ZipA (ΔzipA) or FtsK (ΔftsK) (32, 33, 35). Additionally, FtsAI143L(36) and FtsAE124A(34) weakly bypassed the requirement of FtsN. These mutations also rescued the FtsLL86F/E87K mutant when overexpressed (37). We thus examined whether similar mutations in FtsA of A. baumannii can bypass the need of Aeg1 by constructing substitution FtsA mutants D124A, V144L and Q285W, which correspond to E124A, I143L, and R286W in its E. coli counterpart, respectively (Fig. 3-figure supplement 1) and tested them in strain Δaeg1(PBAD::Aeg1). Whereas FtsAQ285W cannot detectably suppress Aeg1 depletion, expression of FtsAD124A or FtsAV144L partially resecured the growth of strain Δaeg1(PBAD::Aeg1) in medium without ara (Fig. 3D). Despite multiple attempts, we were unable to detect the expression of these FtsA mutants by immunoblotting using antibody specific to the Flag epitope fused to these alleles (Fig. 3D, lower panels), suggesting that the low activity exhibited by these mutants may result from low protein levels.

The observation that mutations in genes involved in cell division bypass the need of Aeg1 suggests that this protein participate in the same cellular process. We thus analyzed the cellular localization of Aeg1 in A. baumannii by introducing a plasmid that expresses Aeg1-mCherry from the PBAD promoter. Fluorescence microscopy analysis revealed that Aeg1 localized at the septum of dividing cells (Fig. 3E), which validated the genetic evidence for its role in cell division.

Aeg1 interacts with multiple core components of the A. baumannii divisome

Earlier version: To further explore the possibility that Aeg1 is a component of the A. baumannii divisome, we determined its interactions with known cell division proteins using the Cya-based bacterial two-hybrid method (36). Among the tested proteins FtsZ, FtsA, ZipA, FtsK, FtsQ, FtsL, FtsB, FtsW, FtsI, and FtsN, Aeg1 detectably interacts with ZipA, FtsK, FtsL, FtsB, and FtsN. ZipA and FtsN appearing to have higher binding affinity with Aeg1 as indicated by the hydrolysis of X-gal by derivatives of the E. coli reporter strain BTH101 (36) harboring appropriate protein fusions (Fig. 4A). When the interactions were quantitatively determined by measuring galactosidase activity, the binding between FtsN and Aeg1 appeared the strongest among all examined protein pairs (Fig. 4B).

Revised version: To further explore the possibility that Aeg1 is a component of the A. baumannii divisome, we determined its interactions with known cell division proteins using the Cya-based bacterial two-hybrid method (38). Effective complementation between the fused T25 and T18 domains requires their localization to the cytoplasmic compartment. Because the predicted transmembrane topology of Aeg1 is not unequivocal, with one model suggesting an N-terminal transmembrane segment and another proposing a more centrally located membrane-spanning domain. To circumvent potential artifacts arising from the ambiguous topological models of Aeg1, we generated two distinct fusion constructs, tethering the T25 domain to either the N-terminus or C-terminus of the Aeg1 protein. We then assessed the complementation between these Aeg1-T25 fusions and the potential binding partners fused to the T18 domain to evaluate the interaction between Aeg1 and other Fts proteins. Among the tested proteins FtsZ, FtsA, ZipA, FtsK, FtsQ, FtsL, FtsB, FtsW, FtsI, and FtsN, Aeg1 detectably interacts with ZipA, FtsK, FtsL, FtsB, and FtsN. Intriguingly, ZipA and FtsN appearing to have higher binding affinity with Aeg1 whether the T25 reporter domain was fused to the N-terminus or C-terminus of the Aeg1 construct (Fig. 4A-D), as indicated by the hydrolysis of X-gal by derivatives of the E. coli reporter strain BTH101 (38) harboring appropriate protein fusions. When the interactions were quantitatively determined by measuring galactosidase activity, the binding of FtsN or ZipA to Aeg1 appeared the strongest among all examined protein pairs (Fig. 4A-D). This observation suggests that the precise topology of Aeg1 does not significantly impact its ability to engage these binding partners.

Aeg1 interacts with multiple core proteins of the A. baumannii divisome.

Earlier version: A-B. Derivatives of strain E. coli BTH101 containing T18-Aeg1 and T25 fusion of the indicated proteins were spotted onto LB agar containing X-gal. Images were acquired after incubation for 18 h at 37°C (A). Interactions determined by quantitative measurement of β-galactosidase activity (B). Results shown were from three independent experiments. Statistical analysis was performed by Student’s t test.

Revised version: A-D. Derivatives of strain E. coli BTH101 containing T18-Aeg1 and T25 fusion of the indicated proteins were spotted onto LB agar containing X-gal. T25 domain was fused to the N-terminus (A-B) or C-terminus (C-D) of Aeg1. Images were acquired after incubation for 18 h at 37°C (A and C). Interactions determined by quantitative measurement of β-galactosidase activity (B and D). Results shown were from three independent experiments. Statistical analysis was performed by Student’s t test.

Earlier version: C. Aeg1 colocalized with a set of core divisome proteins. GFP fusion of ZipA, FtsK, FtsL, FtsB, or FtsW was expressed in strain Δaeg1(pJL03::mCherry-Aeg1). Cells grown in LB broth without ara for 6 h were induced with IPTG to express the GFP fusions for 4 h prior to being processed for imaging. Red arrows indicate the sites of colocalization. Bar, 5 µm. Images are representative of three independent cultures.

Revised version: E. Aeg1 colocalized with a set of core divisome proteins. GFP fusion of ZipA, FtsL, FtsB, or FtsN was expressed in strain WT(pJL03::mCherry-Aeg1). Saturated cultures of bacterial strains harboring the GFP and mCherry fusions were diluted into fresh LB broth. After the cell density of the subcultures reached until OD600=0.6, IPTG (0.25 mM) and ara (0.25%) were added to induce the expression of fusion proteins for 4 h prior to being processed for imaging. Red arrows indicate the sites of colocalization. Cells in which the proteins did not co-localize were shown by white arrows. Bar, 10 µm. Images are representative of three independent cultures.

Earlier version: We next analyzed the localization of Aeg1 and core divisome proteins during cell division of A. baumannii by taking advantage of the elongated morphology of Aeg1-depleted cells, which may allow better resolution of observation via fluorescence microscopy. To this end, we coexpressed the mCherry-Aeg1 fusion with GFP fusion of ZipA, FtsK, FtsL, FtsB, or FtsN in the Δaeg1 mutant from the ara- and IPTG-inducible promoter, respectively. Cultures of strains derived from Δaeg1(PBAD::mCherry-Aeg1) harboring the GFP fusions were induced by ara for 16 h, saturated bacterial cultures were then diluted into fresh LB broth without ara for 6 h to induce the elongation morphology. IPTG (0.25 mM) and ara (0.25%) were added to induce the expression of fusion proteins for 4 h before samples were processed for microscopic analysis. Consistent with results from the protein interaction experiments using the bacterial two-hybrid assay, Aeg1 colocalized with ZipA, FtsK, FtsL, FtsB, and FtsW (Fig. 4C). Thus, Aeg1 interacts with multiple core cell divisome proteins of A. baumannii.

Revised version: We next analyzed the localization of putative Aeg1-interacting proteins identified in our bacterial two-hybrid experiments. To this end, we coexpressed the mCherry-Aeg1 fusion with GFP fusion of ZipA, FtsL, FtsB, or FtsN in the WT strain from the ara- and IPTG-inducible promoter, respectively. Prior to localization studies, we first sought to examine the functionality of the GFP fusions by individually deleting the chromosomal gene of zipA, ftsB, ftsL, and ftsN in A. baumannii strains expressing the fusion proteins from a plasmid. Our results indicate that each GFP fusion was capable of complementing the growth defects caused by deleting the corresponding gene (Fig. 4-figure supplement 1A). Similarly, we found that expression of the mCherry-Aeg1 fusion allowed the deletion of the aeg1 gene (Fig. 4-figure supplement 1B). In each case, the growth of the strain expressing the fusion protein complementing the chromosomal deletion was comparable to that of the wild-type strain (Fig. 4-figure supplement 1A). These strains were then used in experiments aiming at investigating the localization of the relevant proteins. Saturated cultures of strains derived from WT(PBAD::mCherry-Aeg1) (grown for 16 h) that harbored the GFP fusions were diluted into fresh LB broth. When the cell density of the subcultures reached OD600=0.6, IPTG (0.25 mM) or ara (0.25%) were added to induce the expression of fusion proteins for 4 h and samples were processed for microscopic analysis. Consistent with results from the protein interaction experiments using the bacterial two-hybrid assay, we found that Aeg1 colocalized with ZipA, FtsL, FtsB, and FtsN (Fig. 4E, red arrows), but occasional non-co-localization was also observed (Fig. 4E, white arrows). Thus, Aeg1 interacts with multiple core cell divisome proteins of A. baumannii.

Aeg1 is an essential core cell division protein in A. baumannii

Earlier version: The interaction map between Aeg1 and core divisome proteins prompted us to postulate that the absence of Aeg1 prevented these proteins from localizing to the midcell. To test this hypothesis, we introduced plasmids that direct the expression of GFP fusion of each of the divisome proteins into strain Δaeg1(PBAD::mCherry-Aeg1). We evaluated the protein of Aeg1-mCherry at 2, 4, and 6 h after withdrawing arabinose and found that at the 2 h and 4 h time points mCherry-Aeg1 was still readily detectable (Fig. S4). Importantly, we found that removal of arabinose for 6 h rendered Aeg1-mCherry undetectable in approximately 90% of the cells. We thus used the 6 h inducer depletion to examine the effects of Aeg1 depletion. In cells depleted of Aeg1, all of the examined core division proteins displayed midcell mistargeting, including ZipA, FtsK, FtsB, FtsL, and FtsN (Fig. 5A). These results suggest that Aeg1 is a new member of cell division protein, which plays a role in positioning several other core proteins at the septum of A. baumannii.

Revised version: The interaction map between Aeg1 and core divisome proteins prompted us to postulate that the absence of Aeg1 prevented these proteins from localizing to the midcell. To test this hypothesis, we introduced plasmids that direct the expression of GFP fusion of each of the divisome proteins into strain Δaeg1(PBAD::mCherry-Aeg1). We evaluated the protein of Aeg1-mCherry at 2, 4, and 6 h after withdrawing arabinose and found that at the 2 h and 4 h time points mCherry-Aeg1 was still readily detectable (Fig. 5-figure supplement 1). Importantly, we found that removal of arabinose for 6 h rendered Aeg1-mCherry undetectable in approximately 90% of the cells. We thus used the 6 h inducer withdraw to examine the effects of Aeg1 depletion. In cells depleted of Aeg1, all of the examined core division proteins displayed midcell mistargeting, including ZipA, FtsB, FtsL, and FtsN (Fig. 5A). These results suggest that Aeg1 is a new member of cell division protein, which plays a role in positioning several other core proteins at the septum of A. baumannii.

Aeg1 dictates the cellular localization of core divisome proteins.

A. Aeg1 deficiency prevented core divisome proteins from localizing to the middle cell. Cells of derivatives of strain Δaeg1(pJL03::mCherry-Aeg1) expressing GFP fusion of each of the divisome proteins from the PTAC promoter grown in medium with ara were processed for imaging. The red arrows representative examples of the inability of fusion proteins to target to the septal ring. Bar, 5 µm. Images are representative of three parallel cultures.

B-D. A few dominant active mutants of FtsB and FtsW bypassed the need of Aeg1 by A. baumannii. Mutant FtsBE65A (B), FtsBD68H (B), FtsLQ70K (C), FtsWM254I (D), or FtsWS274G (D) was expressed from the IPTG-inducible promoter PTAC in strain Δaeg1(pJL03::Aeg1). Serially diluted cells were spotted onto LB agar containing ara or IPTG. Images were acquired after incubation at 37°C for 18 h (B-D). Similar results were obtained in three independent experiments. The expression of the protein was probed with the HA-specific antibody after IPTG induction (B-D). ICDH was probed as a loading control.

Earlier version: During the formation of the divisome, FtsN was believed to be the last protein being recruited to the septum where it converts FtsQLB and FtsA into their active forms, which subsequently activate the septal peptidoglycan synthase FtsWI (34, 37). If the interaction between Aeg1 and FtsN is critical for the ability of the latter to switch on FtsQLB and FtsA, constitutively active mutants of FtsQLB or Fts A should suppress the need of Aeg1 for cell viability. To test this hypothesis, we expressed HA-FtsBE65A and HA-FtsLQ70K of A. baumannii (equivalent to mutants of their E. coli counterparts that allow the bypass of the requirement of FtsN) (Fig. S3) (34, 38) in strain Δaeg1(PBAD::Aeg1) on medium without ara. Although the expression of HA-FtsBE65A was not detectable with the HA-specific antibody, induction of this mutant completely restored the growth of the Δaeg1 mutant in the absence of Aeg1 (Fig. 5B). In contrast, a similarly constructed plasmid that directs the expression of HA-FtsBD68H cannot bypass the requirement of Aeg1 under the same conditions (Fig. 5B). Similarly expressed HA-FtsLQ70K detectably rescued bacterial growth without Aeg1 at a level considerably lower than that of FtsBE65A (Fig. 5C). The low or undetectable activity exhibited by FtsLQ70K and FtsBD68H may result from the low protein levels in strain Δaeg1(PBAD::Aeg1). Nevertheless, these results support our hypothesis that Aeg1 functions upstream of FtsN, probably by functioning as its activator or recruiting factor.

Revised version: During the formation of the divisome, FtsN was believed to be the last protein being recruited to the septum where it converts FtsQLB and FtsA into their active forms, which subsequently activate the septal peptidoglycan synthase FtsWI (36, 39). If the interaction between Aeg1 and FtsN is critical for the ability of the latter to switch on FtsQLB and FtsA, constitutively active mutants of FtsQLB or FtsA should suppress the need of Aeg1 for cell viability. To test this hypothesis, we expressed HA-FtsBE65A and HA-FtsLQ70K of A. baumannii (equivalent to mutants of their E. coli counterparts that allow the bypass of the requirement of FtsN) (Fig. 5-figure supplement 2) (36, 40) in strain Δaeg1(PBAD::Aeg1) on medium without ara. Although the expression of HA-FtsBE65A was not detectable with the HA-specific antibody, induction of this mutant completely restored the growth of the Δaeg1 mutant in the absence of Aeg1 (Fig. 5B). In contrast, a similarly constructed plasmid that directs the expression of HA-FtsBD68H cannot bypass the requirement of Aeg1 under the same conditions (Fig. 5B). Similarly expressed HA-FtsLQ70K detectably rescued bacterial growth without Aeg1 at a level considerably lower than that of FtsBE65A (Fig. 5C). The low or undetectable activity exhibited by FtsLQ70K and FtsBD68H may result from the low protein levels in strain Δaeg1(PBAD::Aeg1). Nevertheless, these results support our hypothesis that Aeg1 functions upstream of FtsN, probably by functioning as its activator or recruiting factor. It is known that expression of mutants FtsIK211I, FtsWM269I, and FtsWE289G driven by a strong promoter can bypass the need of FtsN in E. coli (41). The ability of FtsBE65A and FtsAE202K to suppress the growth defect of the Δaeg1 mutant prompted us to examine whether constitutively active alleles of FtsW or FtsI have the same effects. Overexpression of FtsWM254I and FtsWS274G of A. baumannii detectably rescued the growth of the strain in which Aeg1 had been depleted (Fig. 5D). FtsI from A. baumannii does not have a residue corresponding to K211 of its E. coli counterpart (Fig. 5-figure supplement 2), we thus did not test the effect of this mutation. Therefore, our results demonstrate that Aeg1 may collaborate with FtsN to activate at least a part of FtsQLB and FtsA for subsequently activation of FtsWI during cell division. Together with results that the variant FtsAE202K suppressed the growth defect caused by aeg1 depletion (Fig. 3B), these observations support a model in which the Aeg1-FtsN complex is involved in the activation of FtsQLB and FtsA.

Removed in the revised version: It is known that expression of mutants FtsIK211I, FtsWM269I, and FtsWE289G driven by a strong promoter can bypass the need of FtsN in E. coli (39). The ability of FtsBE65A and FtsAE202K to suppress the growth defect of the Δaeg1 mutant prompted us to examine whether constitutively active alleles of FtsW or FtsI have the same effects. Overexpression of FtsWM254I and FtsWS274G of A. baumannii detectably rescued the growth of the strain in which Aeg1 had been depleted (Fig. 5D). FtsI from A. baumannii does not have a residue corresponding to K211 of its E. coli counterpart (Fig. S3), we thus did not test the effect of this mutation. Therefore, our results demonstrate that Aeg1 may collaborate with FtsN to activate at least a part of FtsQLB and FtsA for subsequently activation of FtsWI during cell division. Together with results that the variant FtsAE202K suppressed the growth defect caused by aeg1 depletion (Fig. 3B), these observations support a model in which the Aeg1-FtsN complex is involved in the activation of FtsQLB and FtsA.

Discussion

Earlier version: Powerful In-seq DNA sequencing technologies based on saturated transposon insertion libraries have allowed the identification of genes essential for bacterial viability under specific conditions, including commonly used rich media such as LB broth (22). Using a conditional gene deletion method, we have revealed that from 10 putative essential genes identified by this method, three previously annotated as hypotheticals are essential for A. baumannii viability, suggesting that assignment of essentiality to genes identified by methods of this kind need to be individually verified.

Revised version: Powerful In-seq DNA sequencing technologies based on saturated transposon insertion libraries have allowed the identification of genes essential for bacterial viability under specific conditions, including commonly used rich media such as LB broth (22). Using a conditional gene deletion method, we have revealed that from 10 putative essential genes identified by this method, three previously annotated as hypotheticals are essential for A. baumannii viability, suggesting that assignment of essentiality to genes identified by methods of this kind needs to be individually verified.

Earlier version: Our assignment of aeg1 as an essential gene is consistent with a recent study which revealed its essentiality (termed advA) by transposon insertion sequencing, which also observed cell elongation upon the depletion of this gene (29), Our results provided at least two additional lines of evidence to support the notion that Aeg1 is a core cell division protein in A. baumannii. First, random suppression mutations that bypassed the requirement of Aeg1 mapped to a classic cell division gene and mutations in some cell division genes known to circumvent the need of specific cell division proteins rescued the growth defect caused by Aeg1 depletion. Second, Aeg1 interacts with at least a subset of proteins involved in cell division (Fig. 4A-B).

Revised version: Our assignment of aeg1 as an essential gene is consistent with a recent study which revealed its essentiality (termed advA) by transposon insertion sequencing, which also observed cell elongation upon the depletion of this gene (31). Our results provided at least two additional lines of evidence to support the notion that Aeg1 is a core cell division protein in A. baumannii. First, random suppression mutations that bypassed the requirement of Aeg1 mapped to a classic cell division gene and mutations in some cell division genes known to circumvent the need of specific cell division proteins rescued the growth defect caused by Aeg1 depletion (Fig. 3). Second, Aeg1 interacts with at least a subset of proteins involved in cell division (Fig. 4A-D).

Earlier version: In the process of bacterial cell division, over 30 distinct proteins are assembled into a complex typically referred to as “divisome” (40). Ten of these are normally essential for both cell division and survival, and they are now mainly suggested to be the core of the divisome in various bacteria (40, 41). These essential division proteins are mostly named Fts (filamentation temperature sensitive), which stems from the fact that inactivating one of these Fts proteins causes cells to become filamentous, which eventually failed to survive (40, 42). Among these, the FtsQLB complex not only serves as a scaffold for recruitment of subsequent divisome proteins (43, 44), but also plays a an important role in regulating the activity of the divisome (34, 37). The constriction control domain (CCD) of FtsL was previously defined as a critical component for switching the conformation of FtsLB from an off to an on state in order to activate FtsWI by enabling the interaction between FtsIperi and the AWI (activation of FtsWI) region of FtsL (34, 37, 45). Division obstruction caused by Aeg1 depletion can be effectively circumvented by suppression mutations (FtsBE65A) in its CCD domain and by activation mutations (M254I and S274G) of FtsW (Fig. 5D), suggesting that Aeg1 functions upstream of these proteins.

Revised version: In the process of bacterial cell division, over 30 distinct proteins are assembled into a complex typically referred to as “divisome” (42). Ten of these are normally essential for both cell division and survival, and they are now mainly suggested to be the core of the divisome in various bacteria (42, 43). These essential division proteins are mostly named Fts (filamentation temperature sensitive), which stems from the fact that inactivating one of these fts genes causes cells to become filamentous, which eventually failed to survive (42, 44). Among these, the FtsQLB complex not only serves as a scaffold for recruitment of subsequent divisome proteins (45, 46), but also plays an important role in regulating the activity of the divisome (36, 39). The constriction control domain (CCD) of FtsL was previously defined as a critical component for switching the conformation of FtsLB from an off to an on state in order to activate FtsWI by enabling the interaction between FtsIperi and the AWI (activation of FtsWI) region of FtsL (36, 39, 47). Division obstruction caused by Aeg1 depletion can be effectively circumvented by suppression mutations (FtsBE65A) in its CCD domain and by activation mutations (M254I and S274G) of FtsW (Fig. 5D), suggesting that Aeg1 functions upstream of these proteins.

Earlier version: In addition to FtsQLB-mediated activation of FtsI in the periplasm, the constitutively active form of FtsA (FtsA*) acting on FtsW in the cytoplasm, these two events functions synergistically to activate FtsWI for sPG synthesis (35). FtsA is a member of the actin family (46, 47), which interacts with FtsN via a unique IC domain positioned between the IA and IIA domains (48–50). Little is known about the function of the IA and IIA domains of FtsA. That the E202K mutation in the IIA domain bypasses the requirement of Aeg1 by A. baumannii (Fig. 3B) suggests that Aeg1 and FtsA work together to activate FtsW.

Revised version: In addition to FtsQLB-mediated activation of FtsI in the periplasm, the constitutively active form of FtsA (FtsA*) acting on FtsW in the cytoplasm, these two events functions synergistically to activate FtsWI for sPG synthesis (37). FtsA is a member of the actin family (48, 49), which interacts with FtsN via a unique IC domain positioned between the IA and IIA domains (50–52). Little is known about the function of the IA and IIA domains of FtsA. That the E202K mutation in the IIA domain bypasses the requirement of Aeg1 by A. baumannii (Fig. 3B) suggests that Aeg1 and FtsA work together to activate FtsW.

Earlier version: The maturation of the core divisome is facilitated by the recruitment in a nearly sequential order of the late-division proteins FtsK, FtsQ, FtsL/FtsB, FtsW, FtsI, and FtsN (51–53). Once the latest division protein, FtsN, is recruited, division-specific peptidoglycan synthesis conferred by the FtsWI complex is activated to build the septum (53, 54). More recent studies suggest that by binding to FtsA, a small amount of FtsN is recruited to the division site earlier than previously thought (34, 50, 55) where it induces a conformational switch in both FtsA and the FtsQLB subcomplex on both sides of the membrane (34). Furthermore, binding between FtsA and FtsN is regulated by ZipA (50). Our observation of strong binding of Aeg1 to ZipA and FtsN (Fig. 4A-B) prompted us to propose a model in which Aeg1 participates in cell division by anchoring to the division site together with ZipA where it recruits FtsN, which in turn signals to both the cytoplasmic or periplasmic branches of peptidoglycan synthesis involving FtsA and the FtsQLB complex, respectively (34) (Fig. 6A). The placement of Aeg1 as one of the early arrivals in the hierarchical sequential assembly of the divisome is consistent with the fact that the requirement of Aeg1 can be bypassed by constitutively active mutants of FtsAE202K (Fig. 3), FtsB (FtsBE65A), FtsL (FtsLQ70K) and of FtsW (FtsWM254I and FtsWS274G) (Fig. 5).

Revised version: The maturation of the core divisome is facilitated by the recruitment in a nearly sequential order of the late-division proteins FtsK, FtsQ, FtsL/FtsB, FtsW, FtsI, and FtsN (53–55). Once the latest division protein, FtsN, is recruited, division-specific peptidoglycan synthesis conferred by the FtsWI complex is activated to build the septum (55, 56). More recent studies suggest that by binding to FtsA, a small amount of FtsN is recruited to the division site earlier than previously thought (36, 52, 57) where it induces a conformational switch in both FtsA and the FtsQLB subcomplex on both sides of the membrane (36). Furthermore, binding between FtsA and FtsN is regulated by ZipA (52). Our observation of strong binding of Aeg1 to ZipA and FtsN (Fig. 4A-D) prompted us to propose a model in which Aeg1 participates in cell division by anchoring to the division site together with ZipA where it recruits FtsN, which in turn signals to both the cytoplasmic or periplasmic branches of peptidoglycan synthesis involving FtsA and the FtsQLB complex, respectively (36) (Fig. 6A). The placement of Aeg1 as one of the early arrivals in the hierarchical sequential assembly of the divisome is consistent with the fact that the requirement of Aeg1 can be bypassed by constitutively active mutants of FtsAE202K (Fig. 3), FtsB (FtsBE65A), FtsL (FtsLQ70K) and of FtsW (FtsWM254I and FtsWS274G) (Fig. 5).

A model for the function of Aeg1 in cell division of A. baumannii.

Earlier version: In the early stages of Z-ring assembly, anchor proteins FtsA and ZipA tether FtsZ to the inner face of the plasma membrane. In the initial phase of cell division, a small amount of FtsN is recruited to the divisome in a FtsA-dependent manner (left, recruitment phase) (64). The Aeg1-ZipA-FtsN complex acts as a dynamic scaffold for recruiting downstream of Fts proteins. Once more FtsN is recruited to the divisome, it will activate both FtsA and the FtsQLB complex, which in turn activates FtsI in the periplasm (38). Note that constitutively active form of FtsA (FtsA*) acts on FtsW in the cytoplasm (35) so did the hyperactive mutants FtsAE202K. These mutants all bypassed the need of Aeg1 (right, activation phase). Constitutively active mutants were indicated by a star. The red arrows next to FtsAE202K or FtsBE65A indicate overexpression of these mutants.

Revised version: In the early stages of Z-ring assembly, anchor proteins FtsA and ZipA tether FtsZ to the inner face of the plasma membrane. In the initial phase of cell division, a small amount of FtsN is recruited to the divisome in a FtsA-dependent manner (left, recruitment phase) (65). The Aeg1-ZipA-FtsN complex acts as a dynamic scaffold for recruiting downstream of Fts proteins. Once more FtsN is recruited to the divisome, it will activate both FtsA and the FtsQLB complex, which in turn activates FtsI in the periplasm (40). Note that constitutively active form of FtsA (FtsA*) acts on FtsW in the cytoplasm (37) so did the hyperactive mutants FtsAE202K. These mutants all bypassed the need of Aeg1 (right, activation phase). Constitutively active mutants were indicated by a star. The red arrows next to FtsAE202K or FtsBE65A indicate overexpression of these mutants.

Earlier version: It is worth noting that that although homologs of Aeg1 are present in many Gram-negative bacteria, including some recent isolate of E. coli, it appears to be absent in more extensively studied model strains such as K12 (Fig. S2). It is possible that in strain K12, its activity is conferred by a yet unidentified protein that employs a conserved mechanism in cell division. Future study aiming at elucidating the biochemical and structural basis of the activity of Aeg1 and its homologs, and their relationship with other components of the divisome may provide not only insights into the mechanism of bacterial cell division but also leads for the design of antibiotics with novel mechanism of action.

Revised version: It is worth noting that that although homologs of Aeg1 are present in many Gram-negative bacteria, including some recent isolate of E. coli, it appears to be absent in more extensively studied model strains such as K12 (Figure 1-figure supplement 2). It is possible that in strain K12, its activity is conferred by a yet unidentified protein that employs a conserved mechanism in cell division. Future study aiming at elucidating the biochemical and structural basis of the activity of Aeg1 and its homologs, and their relationship with other components of the divisome may provide not only insights into the mechanism of bacterial cell division but also leads for the design of antibiotics with novel mechanism of action.

Materials and methods

Bacterial strains, growth conditions, and plasmids construction

Earlier version: All bacterial strains and plasmids used in this study are listed in Table S3. All A. baumannii strains were derivatives of strain ATCC17978 (56). Bacteria were grown in Luria Bertani (LB) broth or agar. Antibiotics were added to media when necessary. For E. coli, the concentrations used were: ampicillin, 100 µg/mL, kanamycin, 30 μg/mL, chloramphenicol, 30 µg/mL, and gentamicin, 10 μg/mL. For A. baumannii, kanamycin, 30 μg/mL, gentamicin, 10 μg/mL, and streptomycin, 100 μg/mL.

Revised version: All bacterial strains and plasmids used in this study are listed in Supplementary Files 3. All A. baumannii strains were derivatives of strain ATCC17978 (58). Bacteria were grown in Luria Bertani (LB) broth or agar. Antibiotics were added to media when necessary. For E. coli, the concentrations used were: ampicillin, 100 µg/mL, kanamycin, 30 μg/mL, chloramphenicol, 30 µg/mL, and gentamicin, 10 μg/mL. For A. baumannii, kanamycin, 30 μg/mL, gentamicin, 10 μg/mL, and streptomycin, 100 μg/mL.

Earlier version: To construct pUT18C-Flag, two oligos with sequences coding for the Flag tag and appropriate restriction sites were mixed at a 1: 1 ratio, heated to 100°C for 5 min and allowed the mixed DNA to anneal by gradually cooling down to room temperature. The resultant DNA fragment was then inserted into SalI/ClaI digested pUT18C (36). We introduced an EcoRV site into the inserted DNA fragment to facilitate the identification of plasmids carrying the intended insert. To make pJL05 that allows PTAC-driven expression of HA-tagged protein, primers containing the coding sequence of the HA tag were used to amplified the promoter region of pJL02 (24), PCR product was inserted as a PciI/BamHI fragment. This manipulation replaced the Flag tag on pJL02 with an HA tag. pJL03 is a pVRL1-based plasmid harboring PBAD, which allows expressing Flag-tagged proteins in A. baumannii (24).

Revised version: To construct pUT18C-Flag, two oligos with sequences coding for the Flag tag and appropriate restriction sites were mixed at a 1: 1 ratio, heated to 100°C for 5 min and allowed the mixed DNA to anneal by gradually cooling down to room temperature. The resultant DNA fragment was then inserted into SalI/ClaI digested pUT18C (38). We introduced an EcoRV site into the inserted DNA fragment to facilitate the identification of plasmids carrying the intended insert. To make pJL05 that allows PTAC-driven expression of HA-tagged protein, primers containing the coding sequence of the HA tag were used to amplified the promoter region of pJL02 (24), PCR product was inserted as a PciI/BamHI fragment. This manipulation replaced the Flag tag on pJL02 with an HA tag. pJL03 is a pVRL1-based plasmid harboring PBAD, which allows expressing Flag-tagged proteins in A. baumannii (24).

Earlier version: The DNA fragment were amplified from genomic DNA of A. baumannii strain and digested with appropriate restriction enzyme (Table S3), following insert into similarly digested plasmids. The experimental procedure entails fusion of the GFP gene with the target gene, followed by ligation of the resulting construct into the pJL05 vector for subsequent expression.

Revised version: The DNA fragment were amplified from genomic DNA of A. baumannii strain and digested with appropriate restriction enzyme (Supplementary Files 3), following insert into similarly digested plasmids. The experimental procedure entails fusion of the GFP gene with the target gene, followed by ligation of the resulting construct into the pJL05 vector for subsequent expression.

Introduction of plasmids into A. baumannii

Earlier version: Plasmids were introduced into A. baumannii by electroporation (57) or by triparental mating. Electrocompetent cells were prepared using a previously described protocol (58) with minor modifications. Briefly, saturated cultures of A. baumannii were diluted 1:100 into 50 mL LB medium to grow for 18 h at 37°C and cells were harvested by centrifugation (4000g) at room temperature (RT) for 5 min. Cells were washed twice with 10 mL of 10% glycerol, resuspended in 1 mL of 10% glycerol, cells aliquoted in 50 μL. were stored at −80°C. For transformation, 1 μg of plasmid DNA was mixed with 50 μL competent cells and electroporation was performed in 0.2-cm cuvettes on a Bio-Rad Gene Pulser system (2.5 kV/cm, 25 μF, 200 Ω).

Revised version: Plasmids were introduced into A. baumannii by electroporation (59) or by triparental mating. Electrocompetent cells were prepared using a previously described protocol (60) with minor modifications. Briefly, saturated cultures of A. baumannii were diluted 1:100 into 50 mL LB medium to grow for 18 h at 37°C and cells were harvested by centrifugation (4000g) at room temperature (RT) for 5 min. Cells were washed twice with 10 mL of 10% glycerol, resuspended in 1 mL of 10% glycerol, cells aliquoted in 50 μL. were stored at -80°C. For transformation, 1 μg of plasmid DNA was mixed with 50 μL competent cells and electroporation was performed in 0.2-cm cuvettes on a Bio-Rad Gene Pulser system (2.5 kV/cm, 25 μF, 200 Ω).

Earlier version: The E. coli MT607 harboring pRK600 (59) was used for triparental mating. Briefly, 100 μL of overnight cultures of the donor, recipient, and the helper strain were mixed and cells were harvest by centrifugation (4000g) for 5 min. Mixed cells washed twice with 1 mL of sterile water were spotted onto LB agar and the mating was allowed to proceeded for 4-6 h at 37°C before plating onto selective medium.

Revised version: The E. coli MT607 harboring pRK600 (61) was used for triparental mating. Briefly, 100 μL of overnight cultures of the donor, recipient, and the helper strain were mixed and cells were harvest by centrifugation (4000g) for 5 min. Mixed cells washed twice with 1 mL of sterile water were spotted onto LB agar and the mating was allowed to proceeded for 4-6 h at 37°C before plating onto selective medium.

Construction of in-frame deletion mutants of A. baumannii

Earlier version: To delete the putative essential genes, we first introduced a plasmid expressing the gene of interest into A. baumannii and the resulting strains were used to construct chromosomal deletion mutants using a method based on the suicide plasmid pSR47s (25). Briefly, the coding region of the gene was inserted into pJL03, which allows expression of Flag-tagged protein from the PBAD promoter (24) and the resulting plasmid was introduced into A. baumannii by electroporation.

Revised version: To delete the putative essential genes, we first introduced a plasmid expressing the gene of interest into A. baumannii and the resulting strains were used to construct chromosomal deletion mutants using a method based on the suicide plasmid pSR47s (25). Briefly, the coding region of the gene was inserted into pJL03, which allows expression of Flag-tagged protein from the PBAD promoter (24) and the resulting plasmid was introduced into A. baumannii by electroporation. Strains in which the chromosomal deletion is complemented by GFP or mCherry fusions of the relevant Fts proteins were similarly constructed.

Earlier version: To construct the plasmid used for in-frame deletion of A. baumannii genes, 1.0 kilobase DNA fragment was amplified from the upstream and downstream of the gene to delete, respectively, digested with appropriate restriction enzymes and inserted into similarly digested pSR47s (25). The primers were designed so that in each case the gene to be deleted was replaced with an open reading frame consisting of the first and last 15 residues linked by two amino acids encoded by the sequence of the restriction enzyme used to fuse the two DNA fragments (Table S3). The plasmid was then introduced into the A. baumannii strain harboring the pJL03 derivative expressing the gene to be deleted from the PBAD promoter by triparental mating. Transconjugants obtained on LB agar containing 1% ara were grown in LB broth containing 5% sucrose for 18 h prior to being streaked onto LB agar supplemented with 5% sucrose. Deletion mutants were identified by colony PCR with primers corresponding to the 5’ end of the upstream fragment and the 3’ of the downstream fragments.

Revised version: To construct the plasmid used for in-frame deletion of A. baumannii genes, 1.0 kilobase DNA fragment was amplified from the upstream and downstream of the gene to delete, respectively, digested with appropriate restriction enzymes and inserted into similarly digested pSR47s (25). The primers were designed so that in each case the gene to be deleted was replaced with an open reading frame consisting of the first and last 15 residues linked by two amino acids encoded by the sequence of the restriction enzyme used to fuse the two DNA fragments (Supplementary Files 3). The plasmid was then introduced into the A. baumannii strain harboring the pJL03 derivative expressing the gene to be deleted from the PBAD promoter by triparental mating. Transconjugants obtained on LB agar containing 1% ara were grown in LB broth containing 5% sucrose for 18 h prior to being streaked onto LB agar supplemented with 5% sucrose. Deletion mutants were identified by colony PCR with primers corresponding to the 5’ end of the upstream fragment and the 3’ of the downstream fragments.

Ethyl methanesulfonate (EMS) mutagenesis

The A. baumannii strain Δaeg1(PBAD::Aeg1) was grown in LB broth containing 1% ara to saturation, cells from 2-mL culture collected by centrifugation were washed twice with cold buffer A [60 mM K2HPO4, 33 mM KH2PO4, 7.6 mM (NH4)2SO4, 1.7 mM sodium citrate, PH=7]. Cells resuspended in 2x buffer A supplemented with 1.4% EMS were incubated in a shaker (220 rpm) at 37°C for 15 min for 30 min, 45 min and 60 min, respectively. Treated cells harvested by centrifugation were washed twice with buffer A prior to being plated onto LB agar supplemented with 5% glucose. Colonies appeared were purified and tested by spotting diluted cells onto LB agar with or without ara, respectively to confirm the ability to grown in the absence Aeg1. To eliminate the possibility that the ara-independent growth phenotype was due to mutations in the PBAD promoter, the plasmid from each mutant candidate was rescued and the promoter region was sequenced. Only mutants harboring an intact PBAD promoter were retained for further study.

Genome sequencing was performed by Suzhou Golden Wisdom Biological Technology Co., Ltd, and mutations were identified using the genome of the parent strain ATCC 17978 as reference (26).

Bacterial Two-Hybrid assays

Earlier version: To assess interactions between Aeg1 and A. baumannii Fts proteins, we employed the Bordetella pertussis adenylate cyclase (Cya)-based bacterial two-hybrid system (36). The coding region of Aeg1 was fused to the T25 fragment of Cya and testing Fts proteins were individually fused to the T18 fragment on pUT18C-Flag, which allows the detection of the fusion protein with the Flag-specific antibody (Table S3). Interaction was evaluated in the cya- E. coli strain BTH101 (36) by the ability of the testing strains to hydrolyze X-Gal. Quantitative measurement of β-galactosidase activity with performed with ortho-nitrophenyl-β-galactoside (ONPG) using an established protocol (60). Testing bacterial strains were cultured in 2 mL LB broth containing 0.5 mM IPTG for 18 h, and 0.2 mL cells were harvest and resuspended in 1 mL Z-buffer (60 mM Na2HPO4, 40 mM NaH2PO4·H2O, 1 mM MgSO4, 10 mM KCl, and 7 mL of 2-mercaptoethanol per liter). After adding 50 µL 1% (m/v) SDS and 100 μL of chloroform, cells were lysed by 30 s vortex and 50 μL of ONPG (40 mg/mL) was added. The reaction was terminated with 500 µL of 1 M Na2CO3 after ONPG hydrolysis had occurred as determined the presence of the yellow color product. β-galactosidase activity is defined as in Miller units using the following equation: (OD420×1000)/ [(time of incubation in min)x(volume of cell suspension in mL)x(cell density of the culture in OD600)]. All assays were done in triplicate.

Revised version: To assess interactions between Aeg1 and A. baumannii Fts proteins, we employed the Bordetella pertussis adenylate cyclase (Cya)-based bacterial two-hybrid system (38). The coding region of Aeg1 was fused to the T25 fragment of Cya and testing Fts proteins were individually fused to the T18 fragment on pUT18C-Flag, which allows the detection of the fusion protein with the Flag-specific antibody (Supplementary Files 3). Interaction was evaluated in the cya- E. coli strain BTH101 (38) by the ability of the testing strains to hydrolyze X-Gal. Quantitative measurement of β-galactosidase activity with performed with ortho-nitrophenyl-β-galactoside (ONPG) using an established protocol (62). Testing bacterial strains were cultured in 2 mL LB broth containing 0.5 mM IPTG for 18 h, and 0.2 mL cells were harvest and resuspended in 1 mL Z-buffer (60 mM Na2HPO4, 40 mM NaH2PO4·H2O, 1 mM MgSO4, 10 mM KCl, and 7 mL of 2-mercaptoethanol per liter). After adding 50 µL 1% (m/v) SDS and 100 μL of chloroform, cells were lysed by 30 s vortex and 50 μL of ONPG (40 mg/mL) was added. The reaction was terminated with 500 µL of 1 M Na2CO3 after ONPG hydrolysis had occurred as determined the presence of the yellow color product. β-galactosidase activity is defined as in Miller units using the following equation: (OD420×1000)/ [(time of incubation in min)x(volume of cell suspension in mL)x(cell density of the culture in OD600)]. All assays were done in triplicate.

Antibodies and immunoblotting

HA-specific antibody (1:1000) was purchased from Santa Cruz (Cat #sc-7392), and rabbit anti-ICDH (1:2000) was described previously (24). To prepare protein samples, we diluted (1:100) saturated bacterial cultures into fresh LB broth. When OD600 of the cultures reached 0.6, IPTG was added to a final concentration of 0.5 mM. Approximately 1×109 cells were collected, resuspended in 100 μL SDS sample buffer and were boiled for 5 min. Proteins separated by SDS-PAGE were transferred onto nitrocellulose membranes, which were first blocked with 5% nonfat milk for 1 h followed by incubation with appropriate primary antibodies for 14 h. Washed membranes were incubated with the appropriate IRDye-labeled secondary antibodies and signals were detected using an Odyssey CLx system (LI-COR).

Fluorescence imaging

Earlier version: To determine the cellular localization of Aeg1 and its colocalization with members of the divisome, we constructed pJL03::Aeg1-mCherry, which expressed the fusion from the PBAD promoter (23). GFP chimera of Fts proteins were expressed from pJL05 driven by the PTAC promoter inducible by IPTG. Each of the GFP-Fts fusion was introduced into strain Δaeg1(PBAD::mCherry-Aeg1). All strains were maintained in media containing 1% ara to ensure their viability.

Revised version: To determine the cellular localization of Aeg1 and its colocalization with members of the divisome, we constructed pJL03::Aeg1-mCherry, which expressed the fusion from the PBAD promoter (23). GFP chimera of Fts proteins were expressed from pJL05 driven by the PTAC promoter inducible by IPTG. Each of the GFP-Fts fusion was introduced into strain WT(PBAD::mCherry-Aeg1).

To visualize the localization of the fusions, bacterial strains were grown to saturation in LB broth containing ara and the cells washed twice with PBS were diluted to a density of approximately 5×106 cells/mL. The cultures were split into two 2-mL subcultures, the one used to determine how the absence of Aeg1 impacted the cellular localization of Fts proteins received only IPTG but not ara. Cultures were incubated in a shaker for an additional 6 h, and 100 μL cells were withdrawn and fixed by incubation in glutaraldehyde for 20 min at RT. After washing twice with PBS, samples were stained with Hoechst 33342 (1:5000) (Invitrogen, cat#H3570) and mounted on slides for observation with an Olympus IX-83 fluorescence microscope.

Bioinformatics analysis

Earlier version: Aeg1 homologs were identified using tBLASTn (61) using the Aeg1 protein (A1S_3387) from strain ATCC17978 (GenBank: CP000521.1) as the query sequence. Clustal omega (https://www.ebi.ac.uk/Tools/msa/clustalo/) was used to perform multiple sequence alignments. The homology comparison of Fts proteins between E. coli and A. baumannii was also performed using Clustal omega. Jalview was used for editing and viewing sequence alignments (62). The transmembrane regions of Aeg1 were predicted by the SMART Web site (63).

Revised version: Aeg1 homologs were identified using tBLASTn (63) using the Aeg1 protein (A1S_3387) from strain ATCC17978 (GenBank: CP000521.1) as the query sequence. Clustal omega (https://www.ebi.ac.uk/Tools/msa/clustalo/) was used to perform multiple sequence alignments. The homology comparison of Fts proteins between E. coli and A. baumannii was also performed using Clustal omega. Jalview was used for editing and viewing sequence alignments (64). The transmembrane regions of Aeg1 were predicted by the SMART Web sites (28) and Phobius web server (29, 30).

Data quantitation, statistical analyses

Student’s t-test was used to compare the mean between two groups each with at least three independent samples.